<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Instructions" data-toc-modified-id="Instructions-1">Instructions</a></span></li><li><span><a href="#Define-a-new-class" data-toc-modified-id="Define-a-new-class-2">Define a new class</a></span></li><li><span><a href="#Create-a-namedtuple-for-data-modeling" data-toc-modified-id="Create-a-namedtuple-for-data-modeling-3">Create a <code>namedtuple</code> for data modeling</a></span></li><li><span><a href="#Make-a-dataclass--for-server-log-data" data-toc-modified-id="Make-a-dataclass--for-server-log-data-4">Make a <code>dataclass</code>  for server log data</a></span></li><li><span><a href="#Extend-list-class" data-toc-modified-id="Extend-list-class-5">Extend <code>list</code> class</a></span></li></ul></div>

Instructions
--------

- __First thing__ - Change the name of this file to avoid a [merge conflict](https://docs.github.com/en/github/getting-started-with-github/github-glossary#merge-conflict) with GitHub.
- You should type every character for the following activities. 
- Work together with your peers in the same shared code environment. The following environments allow the importation of Jupyter Notebook from GitHub urls:
    - [Google Colab](https://colab.research.google.com)
    - [Deepnote](https://deepnote.com/)
- These activities are not graded.
- They are open resource - feel free to reference documentation, books, videos, and any other resource.
- Attempt each problem. If stuck, move onto next one and come back to a difficult problem.
- During class:
    - Brian will bounce between breakout rooms. Feel free to ask questions or just ignore him.
    - After your group has finished, leave your breakout room and return to the main room.

<!-- grade_id: cell-a0410474f2612710, points: 0, grade: false, solution: true -->
```python
reset -fs
```

Define a new class
------

Make a new class called `Cryptocurrency`

`Cryptocurrency` will have the following attributes:

- `ico` - Initial Coin Offering
- `symbol` - An emoji to represent the coin for easy tagging on Instagram
- `protocol` - The software definition of the network

`Cryptocurrency` will have the following method:

- `check_validity` - Given a block, if the block is FizzBuzzy then the block is valid.

```python
# Solutions

class Cryptocurrency:
    "Model a digitial asset."
    
    def __init__(self, ico:int, symbol:str, protocol:str):
        self.ico      = ico
        self.symbol   = symbol
        self.protocol = protocol
        
    def check_validity(self, block: int):
        "Valid if block is modulo 15"
        return block % 15 == 0 
```

```python
# Tests for Cryptocurrency

msds_coin = Cryptocurrency(ico=10_000_000,
                           symbol="💰",
                           protocol="ethereum")

assert msds_coin.check_validity(block=15) 
assert not msds_coin.check_validity(block=3)
assert not msds_coin.check_validity(block=5)
assert not msds_coin.check_validity(block=14) 
assert not msds_coin.check_validity(block=16) 
```

Create a `namedtuple` for data modeling
----

Imagine you are working for an on-demand car service and need to model pickups.

1. Import `namedtuple` 
1. Create a `namedtuple` to model pick-ups with following attributes
    - driver_id
    - requester_id
    - latitude
    - longitude
    - datestamp

```python
# Solutions
from collections import namedtuple

Pickup = namedtuple('Pickup', ['driver_id', 'requester_id', 'datestamp', 'latitude', 'longitude'])
```

```python
# Test namedtuple

pickup = Pickup(driver_id=33716002811,
                requester_id=3185567453,
                datestamp='01/Aug/2020:13:02:57 -0800',
                latitude=37.773972,
                longitude=-122.431297)

assert pickup[0] == 33716002811
assert {a for a in dir(pickup) if not a.startswith("_")} =={'count', 'datestamp', 'driver_id', 'index', 'latitude', 'longitude', 'requester_id'}
try:
    pickup.driver_id = 1
except AttributeError:
    print("Successfully immutable")
```

Output:
```
Successfully immutable
```

Make a `dataclass`  for server log data
------

Server log files are a raw, unfiltered look at the traffic to a website. They are the text files stored on a webserver. Every time anything (mostly browsers) request anything (mostly urls) from your server, the server adds a line in the log. The logs also store the errors and the attacks on a webserver.

[Apache](https://httpd.apache.org/) is one of the most common software systems to run HTTP servers.

Examples of real Apache web server logs are [here](http://www.monitorware.com/en/logsamples/apache.php).

1. Import `dataclass` 
2. Define a `dataclass` called `ApacheLog` with following attributes:
    - ip
    - datestamp
    - status_code
3. For each attribute, pick the most appropriate type annotations from built-in types.

<!-- grade_id: cell-d20f2c6bd85416c4, points: 0, grade: false, solution: true -->
```python
# Solutions

from dataclasses import dataclass

@dataclass
class ApacheLog:
    "Model an Apache Server Log"
    
    ip: str
    datestamp: str
    status_code: int

```

<!-- grade_id: cell-f967452b7b148e5e, points: 1, grade: true, solution: false -->
```python
# Test the dataclass was created correctly

al = ApacheLog(ip='192.168.0.1',
              datestamp='07/Mar/2020:16:05:49 -0800',
              status_code=200,
              )

assert isinstance(al, ApacheLog)
assert '__dataclass_fields__' in dir(al)
assert {a for a in dir(al) if not a.startswith("__")} == {'datestamp', 'ip', 'status_code'}
```

Extend `list` class
------

Write a new class `CustomList` that inherits from `list`. `CustomList` will everything that a regular `list` does.

In addition, add a method `len` that will return the number of items in the container.

This is example of adding a fluent interface to a class.

```python
# Solutions

class CustomList(list): 
    "Extend built-in list class with additional funcationality."

    def len(self): 
        "Add len as a method. Wrapper for `len` built-in function."
        return len(self) # This is the global len
    
```

```python
# Test the custom list class

cl = CustomList([1, 2, 3])
assert issubclass(CustomList, list)
assert cl.len() == len(cl)
```

<br>
<br> 
<br>

----

<br>
<br> 
<br>

----